<a href="https://colab.research.google.com/github/Gift-eiei/CE156_CourseWork_Assigment/blob/main/week5/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "PHRUGSA LIMBUNLOM (2311569)"
COLLABORATORS = "PHRUGSA LIMBUNLOM (2311569)"
```

---

---
# Welcome to CE880
### This is your week-4 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

### Task: Complete all the three questions given as follows:

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
# !pip install -U scikit-learn # Run if required
```

```python
import pandas as pd
import numpy as np
from sklearn.cluster import DBSCAN
from sklearn.cluster import KMeans
from sklearn.datasets import make_circles, make_moons
from sklearn.metrics import adjusted_rand_score
from sklearn.preprocessing import StandardScaler
from sklearn.metrics import silhouette_score, adjusted_rand_score
from sklearn.decomposition import PCA
from sklearn import cluster, datasets, metrics
scaler = StandardScaler()
```

---
## Question 1:
We have already loaded `moon` data for you using `make_moons` function.

Write a code to perform `k-means` clustering on `scaled_features` variable and return `silhouette_score` as an output

<!-- grade_id: myScore_kmean_ans, points: 0, grade: false, solution: true -->
```python
# Let's make two moons, where n=250 is the number of samples
features, true_labels = make_moons(n_samples=250, noise=0.05, random_state=42)
# fit_transform() method is basically the combination of fit method and transform method,
# it is equivalent to fit().transform(). This method performs fit and transform on the
# input data at a single time and converts the data points.
# If we use fit and transform separate when we need both then it
# will decrease the efficiency of the model so we use fit_transform()
# which will do both the work.
scaled_features = scaler.fit_transform(features)
myscore = []

def myScore_kmean(scaled_features):
    """Write code to perform k-means clustering and print silhouette_score"""
    # YOUR CODE HERE
    km=cluster.KMeans(n_clusters=2)
    km.fit(scaled_features)
    km.labels_
    myscore = metrics.silhouette_score(scaled_features, km.labels_)
    return myscore
    raise NotImplementedError()
```

<!-- grade_id: myScore_kmean_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(myScore_kmean(scaled_features), 0.5, rel_tol=0.10)
```

Output:
```
/usr/local/lib/python3.10/dist-packages/sklearn/cluster/_kmeans.py:870: FutureWarning: The default value of `n_init` will change from 10 to 'auto' in 1.4. Set the value of `n_init` explicitly to suppress the warning
  warnings.warn(
```

---
## Question 2:

Write a code to perform `DBSCAN` clustering algorithm on `scaled_features` and return `silhouette_score` as an output. Please set the `eps` = 0.2 and `min_samples` = 5

<!-- grade_id: myScore_DBSCAN_ans, points: 0, grade: false, solution: true -->
```python

# Let's make two moons, where n is the number of samples
from sklearn.datasets import make_moons
features, true_labels =  make_moons(n_samples=250, noise=.1)

def myScore_DBSCAN(scaled_features):
    """Write code to perform DBSCAN algorithm and print silhouette_score"""
    # YOUR CODE HERE
    dbs = cluster.DBSCAN(eps=0.2, min_samples=5)
    dbs.fit(scaled_features)
    dbs.labels_
    return metrics.silhouette_score(scaled_features, dbs.labels_)
    raise NotImplementedError()
```

<!-- grade_id: myScore_DBSCAN_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
assert math.isclose(myScore_DBSCAN(scaled_features), 0.2, rel_tol=0.10)
```

## Question 3:

Complete this code perform `PCA`  on `iris data` and return `explained_variance_ratio_` as an output for the first two principal components (PCs)

<!-- grade_id: myPCA_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
url = "https://archive.ics.uci.edu/ml/machine-learning-databases/iris/iris.data"
# load dataset into Pandas DataFrame
df = pd.read_csv(url, names=['sepal length','sepal width','petal length','petal width','target'])

features = ['sepal length', 'sepal width', 'petal length', 'petal width']
# Separating out the features
x = df.loc[:, features].values
# Separating out the target
y = df.loc[:,['target']].values
# Standardizing the features
x = StandardScaler().fit_transform(x)

def myPCA(x):
    """Complete this code return the explained variance ratio for the first two prinical
    components after doing PCA on x """
    # YOUR CODE HERE
    pca = PCA()
    pca.fit_transform(x)
    return pca.explained_variance_ratio_[0:2]
    raise NotImplementedError()
```

<!-- grade_id: myPCA_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
x= myPCA(x)
y = [0.7277045209380135, 0.23030523267680633]
np.testing.assert_allclose(x, y, rtol=1e-5, atol=0)

```